```python
# Initialize Otter
import otter
grader = otter.Notebook("guess_my_word.ipynb")
```

# Guess My Word

#### Authors:
v1.0 (Spring 2020) William Gan, Kannan Ramchandran
v1.1 (Fall 2024) Tianhao Wu

## Introduction

In Guess My Word http://hryanjones.com/guess-my-word/, you make guesses at a secret word and the computer tells you if it's before or after. Being bad at it, Efe wants to write a program to solve it. In this lab, you'll explore a couple ideas he has and help him write the program.

```python
import matplotlib.pyplot as plt
import numpy as np

%matplotlib inline
```

## Modeling the Game

Suppose we have obtained a list of the words used in the game. Each word has a frequency representing relatively how often it appears in the English language. Suppose the secret word is chosen proportionally to this frequency. The code in the cells below simulate the game. A couple things to note:

- The frequency is not a probability i.e. it is unnormalized.
- In this version of the game, we're assuming guesses come from the list of possible words used.
- We're also assuming that you aren't told the secret word if you guess it. You have to deduce it yourself.

```python
WORD_FREQ = {'course': 78, 'against': 221, 'ready': 115, 'daughter': 64, 'work': 106, 'friends': 32, 'minute': 21, 'though': 339, 'supposed': 21, 'honey': 6, 'point': 64, 'start': 20, 'check': 9, 'alone': 154, 'matter': 104, 'office': 8, 'hospital': 18, 'three': 220, 'already': 298, 'anyway': 2, 'important': 110, 'tomorrow': 34, 'almost': 146, 'later': 94, 'found': 212, 'trouble': 23, 'excuse': 21, 'money': 79, 'different': 115, 'between': 233, 'every': 293, 'party': 45, 'either': 94, 'enough': 33, 'year': 53, 'house': 231, 'story': 49, 'crazy': 1, 'mind': 125, 'break': 27, 'tonight': 5, 'person': 54, 'sister': 30, 'pretty': 44, 'trust': 7, 'funny': 10, 'gift': 4, 'change': 69, 'business': 57, 'train': 21, 'under': 360, 'close': 98, 'reason': 87, 'today': 24, 'beautiful': 65, 'brother': 69, 'since': 159, 'bank': 13, 'yourself': 58, 'without': 633, 'until': 18, 'forget': 62, 'anyone': 112, 'promise': 22, 'happy': 141, 'worry': 3, 'school': 1, 'afraid': 116, 'cause': 111, 'doctor': 80, 'exactly': 23, 'second': 99, 'look': 359, 'feel': 121, 'somebody': 8, 'stuff': 1, 'morning': 121, 'heard': 411, 'world': 84, 'chance': 41, 'call': 74, 'watch': 19, 'whatever': 41, 'perfect': 16, 'dinner': 89, 'family': 80, 'heart': 101, 'least': 77, 'answer': 97, 'woman': 95, 'bring': 71, 'probably': 67, 'question': 144, 'stand': 40, 'truth': 32, 'problem': 13, 'patch': 3, 'pass': 54, 'famous': 19, 'true': 63, 'power': 124, 'cool': 2, 'last': 316, 'fish': 4, 'remote': 12, 'race': 7, 'noon': 4, 'wipe': 3, 'grow': 26, 'learn': 15, 'itself': 117, 'chip': 4, 'print': 2, 'young': 397, 'argue': 3, 'clean': 24, 'remove': 6, 'flew': 43, 'replace': 5, 'side': 218, 'walk': 28, 'gate': 28, 'finger': 20, 'target': 1, 'judge': 25, 'push': 11, 'thought': 608, 'wear': 10, 'desert': 1, 'relief': 14, 'bright': 54, 'deal': 22, 'father': 182, 'machine': 9, 'know': 623, 'step': 49, 'exercise': 14, 'present': 90, 'wing': 9, 'ship': 11, 'wait': 42, 'fancy': 14, 'eight': 31, 'hall': 31, 'rise': 12, 'river': 27, 'round': 351, 'girl': 84, 'winter': 29, 'speed': 14, 'long': 452, 'oldest': 3, 'lock': 10, 'kiss': 30, 'garden': 32, 'fight': 56, 'hook': 2, 'desk': 3, 'test': 6, 'serious': 55, 'exit': 1, 'branch': 5, 'keyboard': 1, 'naked': 18, 'science': 32, 'trade': 2, 'quiet': 57, 'home': 91, 'prison': 2, 'blue': 64, 'window': 64, 'whose': 101, 'spot': 39, 'dark': 74, 'create': 2, 'quick': 23, 'face': 609, 'terrible': 122, 'accept': 24, 'door': 186, 'touch': 29, 'care': 37, 'rescue': 5, 'real': 66, 'title': 7, 'city': 39, 'fast': 18, 'season': 1, 'town': 67, 'picture': 11, 'tower': 4, 'engine': 2, 'lift': 23, 'respect': 34, 'time': 638, 'mission': 8, 'play': 41, 'discover': 9, 'half': 118, 'unusual': 13, 'ball': 66, 'tool': 3, 'heavy': 47, 'night': 124, 'farm': 2, 'firm': 35, 'gone': 133, 'help': 128, 'easy': 38, 'library': 3, 'group': 54, 'taste': 10, 'large': 149, 'imagine': 56, 'normal': 6, 'outside': 39, 'paper': 39, 'nose': 19, 'doing': 103, 'moon': 6, 'hour': 50, 'protect': 9, 'hate': 13, 'dead': 46, 'double': 15, 'nothing': 318, 'restaurant': 2, 'reach': 32, 'note': 27, 'tell': 327, 'baby': 20, 'future': 56, 'tall': 39, 'drop': 16, 'speak': 168, 'rule': 14, 'pair': 19, 'ride': 41, 'game': 30, 'hair': 94, 'hurt': 21, 'allow': 44, 'oven': 2, 'live': 71, 'horse': 154, 'bottle': 31, 'rock': 1, 'public': 28, 'find': 158, 'green': 38, 'heat': 19, 'plan': 62, 'mean': 35, 'little': 513, 'spend': 24, 'nurse': 29, 'practice': 5, 'wish': 165, 'uncle': 2, 'core': 1, 'stop': 40, 'number': 66, 'nest': 2, 'pool': 5, 'message': 12, 'active': 23, 'throw': 25, 'pull': 10, 'level': 11, 'wrist': 1, 'hold': 59, 'huge': 36, 'finish': 26, 'teeth': 21, 'flag': 7, 'head': 382, 'private': 24, 'together': 72, 'child': 51, 'decide': 16, 'listen': 46, 'jealous': 12, 'wide': 24, 'straight': 92, 'fall': 63, 'joke': 16, 'table': 116, 'spread': 58, 'deep': 69, 'quit': 3, 'save': 36, 'worst': 18, 'email': 2, 'glass': 44, 'scale': 1, 'safe': 7, 'path': 48, 'excellent': 26, 'place': 229, 'zone': 1, 'luck': 13, 'sign': 43, 'report': 48, 'myself': 83, 'knee': 12, 'need': 97, 'root': 1, 'light': 96, 'sure': 60, 'page': 10, 'life': 331, 'space': 22, 'magic': 4, 'size': 6, 'food': 22, 'period': 34, 'mistake': 16, 'full': 153, 'paid': 41, 'horrible': 6, 'special': 70, 'hidden': 20, 'rain': 15, 'field': 104, 'kick': 1, 'ground': 57, 'screen': 7, 'risky': 1, 'human': 102, 'nobody': 20, 'high': 95, 'class': 15, 'street': 28, 'cold': 93, 'metal': 6, 'nervous': 12, 'wind': 22, 'summer': 22, 'president': 1, 'empty': 23, 'square': 10, 'popular': 8, 'loud': 40, 'online': 4, 'something': 479, 'knot': 7, 'mark': 8, 'road': 113, 'storm': 9, 'said': 2406, 'floor': 25, 'theater': 11, 'kitchen': 6, 'action': 122, 'equal': 24, 'nice': 24, 'dream': 15, 'sound': 144, 'fifth': 17, 'talk': 199, 'police': 20, 'draw': 24, 'bunch': 1, 'idea': 54, 'jerk': 3, 'copy': 10, 'success': 34, 'favor': 18, 'open': 84, 'neat': 1, 'gold': 28, 'free': 116, 'mile': 20, 'lying': 67, 'meat': 3, 'nine': 19, 'wonderful': 15, 'hero': 17, 'quilt': 5, 'move': 60, 'early': 43, 'remember': 79, 'understand': 255, 'month': 20, 'everyone': 160, 'quarter': 18, 'center': 36, 'universe': 3, 'name': 72, 'inside': 14, 'yell': 4, 'jacket': 10, 'nation': 16, 'support': 15, 'lunch': 7, 'twice': 39, 'boot': 11, 'alive': 16, 'build': 5, 'date': 4, 'room': 385, 'fire': 115, 'music': 27, 'leader': 8, 'rest': 70, 'plant': 5, 'connect': 1, 'land': 15, 'body': 66, 'belong': 12, 'trick': 4, 'wild': 11, 'quality': 8, 'band': 17, 'health': 29, 'love': 345, 'hand': 314, 'yeah': 1, 'dozen': 16, 'glove': 7, 'give': 249, 'thick': 26, 'flow': 12, 'project': 9, 'tight': 11, 'join': 42, 'cost': 24, 'trip': 2, 'lower': 34, 'angry': 86, 'line': 74, 'rich': 34, 'owner': 8, 'block': 5, 'shut': 19, 'neck': 37, 'write': 45, 'hotel': 2, 'danger': 40, 'impossible': 113, 'illegal': 2, 'show': 100, 'come': 493, 'want': 211, 'click': 5, 'none': 34, 'done': 224, 'hope': 65, 'share': 25, 'water': 48, 'dust': 16, 'handle': 8, 'unhappy': 12, 'guess': 7, 'past': 119, 'frame': 14, 'knob': 2, 'ugly': 3, 'lesson': 4, 'bear': 40, 'gross': 3, 'midnight': 8, 'grass': 8, 'middle': 73, 'birthday': 1, 'rose': 147, 'useless': 21, 'hole': 5, 'drive': 45, 'color': 12, 'sell': 4, 'unfair': 1, 'send': 59, 'crash': 2, 'knife': 2, 'wrong': 36, 'guest': 7, 'strong': 51, 'weather': 15, 'undo': 2, 'catch': 23, 'neighbor': 7, 'stream': 3, 'random': 4, 'continue': 16, 'return': 81, 'begin': 52, 'kitten': 6, 'thin': 65, 'pick': 8, 'whole': 522, 'useful': 12, 'rush': 16, 'mine': 11, 'toilet': 2, 'enter': 58, 'wedding': 9, 'wood': 32, 'meet': 101, 'stolen': 7, 'hungry': 14, 'card': 15, 'fair': 9, 'crowd': 152, 'glow': 16, 'ocean': 1, 'peace': 59, 'match': 17, 'hill': 45, 'welcome': 6, 'across': 111, 'drag': 5, 'island': 7, 'edge': 22, 'great': 225, 'feet': 77, 'iron': 3, 'wall': 19, 'fill': 5, 'boat': 3, 'hard': 64, 'happen': 45, 'tiny': 6, 'event': 57, 'recently': 10, 'seven': 52, 'tree': 17, 'rough': 13, 'secret': 33, 'nature': 15, 'short': 87, 'inch': 1, 'raise': 16, 'warm': 32, 'gentle': 30, 'roll': 13, 'search': 20, 'regular': 27, 'here': 206, 'count': 232, 'hunt': 15, 'keep': 75, 'week': 21, 'venerable': 4, 'exemplary': 4, 'extenuating': 1, 'mundane': 2, 'hypothesis': 1, 'scrutinize': 1, 'tenacious': 1, 'prosperity': 2, 'compromise': 3, 'submissive': 7, 'wary': 1, 'amicable': 1, 'superficial': 2, 'brusque': 1, 'reconciliation': 1, 'brazen': 1, 'benevolent': 2, 'conditional': 2, 'disdain': 4, 'asylum': 4, 'compassion': 1, 'condescending': 6, 'querulous': 4, 'inevitable': 21, 'renovation': 1, 'provocative': 2, 'subtle': 15, 'diligent': 2, 'orator': 2, 'superfluous': 6, 'tactful': 2, 'restrained': 8, 'abstain': 2, 'reverence': 1, 'spontaneous': 1, 'haughty': 1, 'parched': 1, 'impute': 1, 'transient': 2, 'confluence': 2, 'conventional': 2, 'sufficient': 14, 'negligence': 2, 'temporary': 3, 'regret': 14, 'words': 265, 'diminish': 2, 'plausible': 1, 'necessity': 42, 'diminutive': 2, 'capture': 32, 'keen': 9, 'delirious': 2, 'lunatic': 5, 'fraternity': 1, 'loathe': 1, 'entrenched': 4, 'foment': 2, 'counteract': 1, 'subsequent': 2, 'eccentric': 1, 'decree': 4, 'devise': 5, 'enumerate': 1, 'caricature': 1, 'ridiculous': 7, 'dispatch': 11, 'sacrilegious': 1, 'saucer': 1, 'undermine': 1, 'superb': 1, 'exhilarating': 1, 'suspicious': 2, 'garment': 2, 'typical': 1, 'coarse': 12, 'repudiate': 1, 'frock': 2, 'theoretical': 1, 'contemplate': 5, 'anticipate': 2, 'confide': 3, 'constitution': 2, 'ferocious': 2, 'justify': 12, 'reinforce': 1, 'virtuous': 10, 'dilemma': 1, 'provincial': 5, 'establish': 3, 'renounce': 8, 'invade': 1, 'coincide': 2, 'deference': 3, 'scorn': 2, 'dignity': 19, 'complacent': 3, 'sheath': 1, 'austere': 2, 'nuisance': 4, 'jarring': 1, 'amiable': 12, 'condemn': 6, 'indifferent': 17, 'stupendous': 1, 'theology': 2, 'confound': 2, 'retention': 1, 'provision': 3, 'knight': 2, 'callous': 1, 'refute': 1, 'restoration': 3, 'momentous': 1, 'reconsider': 1, 'withstand': 5, 'virile': 1, 'calamity': 3, 'hypothetical': 1, 'demonstrate': 2, 'flourish': 2, 'peculiar': 28, 'hostile': 17, 'vicious': 10, 'malicious': 1, 'stark': 1, 'perish': 9, 'intimate': 37, 'allegory': 1, 'righteous': 1, 'reconcile': 3, 'saber': 17, 'placid': 1, 'universal': 23, 'recompense': 1, 'jaded': 1, 'astute': 1, 'concise': 1, 'surly': 1, 'recite': 3, 'quench': 1, 'valiant': 4, 'fatigue': 5, 'remnant': 2, 'query': 1, 'marauder': 2, 'chalice': 1, 'hitherto': 9, 'brute': 2, 'narrative': 2, 'supercilious': 1, 'heritage': 1, 'novice': 1, 'delusion': 3, 'grim': 4, 'quizzical': 1, 'conduct': 13, 'tension': 2, 'expend': 1, 'tremendous': 5, 'providence': 1, 'robust': 2, 'altercation': 1, 'convey': 3, 'substantial': 2, 'sumptuous': 1, 'frankly': 9, 'yield': 9, 'emphatic': 3, 'adhere': 1, 'conceive': 12, 'reform': 7, 'delicate': 27, 'contradict': 10, 'sinister': 2, 'arsenal': 2, 'vital': 7, 'brunt': 2, 'jolt': 1, 'rapture': 6, 'conviction': 20, 'deliberate': 6, 'discord': 2, 'vigor': 4, 'allude': 1, 'envy': 10, 'doctrine': 1, 'cultivate': 1, 'commission': 5, 'reluctant': 9, 'braggart': 1, 'stature': 1, 'eloquent': 2, 'fiscal': 1, 'competent': 1, 'apprehensive': 2, 'transparent': 5, 'horizontal': 1, 'gauge': 1, 'prompt': 3, 'endure': 15, 'subordinate': 6, 'lucrative': 2, 'exploit': 2, 'assertion': 1, 'innate': 2, 'exert': 2, 'election': 1, 'reproach': 21, 'notorious': 5, 'billet': 1, 'dismay': 9, 'imposing': 4, 'transaction': 1, 'bolster': 1, 'stationary': 2, 'irrational': 3, 'dreary': 1, 'novel': 6, 'quaint': 1, 'recalcitrant': 1, 'jovial': 2, 'responsibility': 18, 'pinnacle': 1, 'forage': 4, 'indulge': 2, 'annihilate': 3, 'zenith': 1, 'sentiment': 2, 'fundamental': 3, 'principle': 7}
```

```python
total_freq = sum(WORD_FREQ.values())
sorted_words = sorted(list(WORD_FREQ))
sorted_word_freqs = np.array([WORD_FREQ[w] for w in sorted_words])
cum_freqs = np.cumsum(sorted_word_freqs)
```

```python
print(f'Total frequency: {total_freq}')
print(f'First 10 words sorted: {sorted_words[:10]}')
print(f'Frequencies associated with these words: {sorted_word_freqs[:10]}')
print(f'Cumulative freqs associated with these words: {cum_freqs[:10]}')
```

The following function will use binary search to find the word associated with an idx, which should be from 0 to total_freq-1. For example, `binary_search_word_generation(126)` will return the 3rd word, "across".

```python
def binary_search_word_generation(idx):
    i, j = 0, len(sorted_words) - 1
    while i < j:
        mid = (i + j) // 2
        if cum_freqs[mid] > idx:
            j = mid
        else:
            i = mid + 1
    return i
```

**In the following cell, implement `generate_secret_word` using `binary_search_word_generation`.**

```python
def generate_secret_word():
    """Generate a random word from the word list, weighted by word frequency.
    
    You should approach this problem by:
    1. Generating a random integer between 0 and total_freq-1.
    2. Using binary_search_word_generation to find the word associated with that integer.
    """
    # BEGIN SOLUTION
    rand_idx = np.random.randint(total_freq)
    return sorted_words[binary_search_word_generation(rand_idx)]
    # END SOLUTION

assert generate_secret_word() in WORD_FREQ
```

```python
grader.check("generate_secret_word")
```

This cell creates a game by generating a secret word and returning `guess_function`. `guess_function` returns True if the guess is lexicographically "less than" the secret word.

```python
def create_game(secret_word):
    def guess_function(guess):
        assert guess in WORD_FREQ
        return guess < secret_word
    return guess_function
```

## Binary Search Idea

The first idea is to ignore the frequencies and binary search on the correct word. In the cell below, `binary_search_game` receives a `guess_function` which was created by a call to `create_game`. Your code should use this `guess_function` to deduce what the secret word was, as well as keep track of how many guesses it took. **In the cell below, write code that returns the deduced secret word and number of guesses required**.  

*Hint: If you want help implementing binary search, you could model it off of `binary_search_word_generation`.*

```python
def binary_search_game(guess_function):
    """Use binary search to guess the secret word."""
    guesses = 0
    deduced_word = None
    # BEGIN SOLUTION
    i, j = 0, len(sorted_words) - 1
    while i < j:
        mid = (i + j) // 2
        word_to_guess = sorted_words[mid]
        if guess_function(word_to_guess):
            i = mid + 1
        else:
            j = mid
        guesses += 1
    deduced_word = sorted_words[i]
    # END SOLUTION
    return deduced_word, guesses
```

```python
grader.check("binary_search_game")
```

Use the following function to test your code. **While your implementation may not get the exact same number of guesses, it should take 9 or 10 guesses**. **It should also be deducing the correct word**.

```python
def test_binary_search_game():
    '''Sanity checks for binary_search_game.
    
    While your implementation may give different numbers, the number of guesses
    should be 9 or 10.
    '''
    words = ['find', 'time', 'great', 'anyone', 'lying']
    print([binary_search_game(create_game(w)) for w in words])
```

```python
# Staff solution gets [('find', 10), ('time', 10), ('great', 10), ('anyone', 10), ('lying', 9)].
test_binary_search_game()
```

## Huffman Coding Idea

The second idea uses the frequencies and is similar to [Huffman Coding](https://en.wikipedia.org/wiki/Huffman_coding).

It works as follows:

1. Sort all the words in alphabetical order in a list.
2. Go through all consecutive pairs of words. For each pair, consider the sum of the two frequencies. Find the pair whose sum is minimal.
3. Combine those two words into a node, and make the frequency their sum. The node's left child will be the first word in the pair and the node's right child will be the second word in the pair.
4. Replace the two words with the node in the list.
5. Go back to step 1, and now treating nodes and words as the same, repeat until there's only one node left.
6. This node is the root of the Huffman Tree.

When creating new nodes, we also need to set `rightmost_word`. It'll be useful when we try to guess the secret word, which works as follows.

1. Set the current node to the root.
2. Guess the `rightmost_word` of the current node's left child.
3. If this is before the secret word, than the secret word must be in the right child. Otherwise, it's in the left child.
4. Go back to step 2. Repeat until we've reached a leaf in this tree.

```python
class HuffmanNode:
    def __init__(self, freq, word=None):
        # The total frequency of the words in our own subtree.
        self.freq = freq
        # Only non-null in leaf nodes.
        self.word = word
        # The rightmost word in our own subtree.
        self.rightmost_word = word
        # Our children nodes
        self.left_subtree = None
        self.right_subtree = None
```

**In the following cell, fill code to generate the HuffmanTree.**

```python
def build_huffman_tree():
    sorted_words = sorted(list(WORD_FREQ))
    nodes = [HuffmanNode(WORD_FREQ[w], w) for w in sorted_words]
    # Will need to combine two nodes len(WORD_FREQ) - 1 times.
    for i in range(len(WORD_FREQ) - 1):
        min_node_idx = None
        min_freq = None
        # Choose the consecutive pair of nodes with smallest frequency sum to
        # combine.
        for j in range(0, len(nodes) - 1):
            freq = nodes[j].freq + nodes[j+1].freq
            if min_freq is None or freq < min_freq:
                min_node_idx = j
                min_freq = freq
        new_node = HuffmanNode(min_freq)
        # Setup the new node and then replace the current two nodes in the list
        # with the new node.
        
        # TODO: Implement the rest of the function
        # BEGIN SOLUTION
        new_node.left_subtree = nodes[min_node_idx]
        new_node.right_subtree = nodes[min_node_idx + 1]
        new_node.rightmost_word = new_node.right_subtree.rightmost_word
        nodes.pop(min_node_idx)
        nodes.pop(min_node_idx)
        nodes.insert(min_node_idx, new_node)
        # END SOLUTION
    return nodes[0]
```

```python
grader.check("build_huffman_tree")
```

**In the cell below, write code that plays a game with the HuffmanTree and returns the deduced word and number of guesses.**


```python
root = build_huffman_tree()

def huffman_game(guess_function):
    """Use the Huffman Tree to guess the secret word."""
    guesses = 0
    deduced_word = None
    # BEGIN SOLUTION
    curr_node = root
    while curr_node.left_subtree:
        if guess_function(curr_node.left_subtree.rightmost_word):
            curr_node = curr_node.right_subtree
        else:
            curr_node = curr_node.left_subtree
        guesses += 1
    deduced_word = curr_node.word
    # END SOLUTION
    return deduced_word, guesses
```

```python
grader.check("huffman_game")
```

Use the following function to test your code. **While your implementation may not get the exact same number of guesses, it should be deducing the correct word**.

```python
def test_huffman_game():
    '''Sanity checks for huffman_game.
    
    The number of guesses may not be less than binary search for particular
    examples.
    '''
    words = ['find', 'time', 'great', 'anyone', 'lying']
    print([huffman_game(create_game(w)) for w in words])
```

```python
# Staff solution gets [('find', 8), ('time', 6), ('great', 8), ('anyone', 8), ('lying', 10)].
test_huffman_game()
```

## Analysis

In this section, we'll try to compare the average number of guesses required by each of the two methods and connect it to entropy.

**In the following cell, compute the entropy of the distribution for the secret word.**

*Note: By default, np.log uses the natural log.*

```python
entropy = None
"""Fill out the rest of the code to calculate the entropy of the distribution for the secret word."""
# BEGIN SOLUTION
probs = sorted_word_freqs / total_freq
entropy = np.sum(- probs * np.log2(probs))
# END SOLUTION
print(entropy)
```

We'll use sampling to estimate the mean number of guesses for each method. In particular, if $X_i$ is the number of guesses required for trial $i$, we can approximate
$$
E[X] \approx \frac{1}{n} \sum_{i=1}^n X_i
$$

```python
SAMPLES = 100000
binary_search = np.zeros(SAMPLES)
huffman = np.zeros(SAMPLES)
for i in range(SAMPLES):
    secret_word = generate_secret_word()
    _, binary_search[i] = binary_search_game(create_game(secret_word))
    _, huffman[i] = huffman_game(create_game(secret_word))
```

```python
bins = np.arange(4, 15)
plt.hist(binary_search, bins, alpha=0.5, label='Binary Search')
plt.hist(huffman, bins, alpha=0.5, label='Huffman')
plt.xlabel('Number of Guesses')
plt.legend()
plt.title('Guess My Word')
```

Let's create confidence intervals for $E[X]$. We know that $\frac{1}{n} \sum_{i=1}^n X_i \sim \mathcal{N}(E[X], \frac{var(X)}{n})$ by the CLT. **In the following cell, calculate `sigma_n` and `mu`**. We can use the sample mean to estimate $E[X]$ and sample variance to estimate $var(X)$.

```python
def confidence_interval_95(samples):
    mu, sigma_n = 0, 0
    # BEGIN SOLUTION
    sigma_n = np.std(samples) / np.sqrt(len(samples))
    mu = np.mean(samples)
    # END SOLUTION
    return f'{mu:.3f} +- {1.96 * sigma_n:.3f}'
```

```python
grader.check("confidence_interval_95")
```

```python
print(f'Binary Search: {confidence_interval_95(binary_search)}')
print(f'Huffman: {confidence_interval_95(huffman)}')
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

We recommend you to direct submit the .ipynb file after you finish the lab.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for binary_search_game (points: 2)
```
>>> def test_binary_search_game_public():
...     """tests for binary_search_game."""
...     test_word = 'test'
...     (deduced, guesses) = binary_search_game(create_game(test_word))
...     assert deduced == test_word, f'Expected {test_word} but got {deduced}'
...     assert isinstance(guesses, int), f'Expected guesses to be int but got {type(guesses)}'
...     assert guesses > 0, f'Expected positive number of guesses but got {guesses}'
...     assert guesses <= 12, f'Used more guesses than words'
...     for word in ['find', 'time', 'great', 'anyone', 'lying']:
...         (deduced, guesses) = binary_search_game(create_game(word))
...         assert deduced == word, f'Expected {word} but got {deduced}'
...         assert isinstance(guesses, int)
...         assert guesses > 0
...         assert guesses <= 12
>>> test_binary_search_game_public()
```


Tests for build_huffman_tree (points: 2)
```
>>> def collect_words(node, words):
...     if node.word:
...         words.add(node.word)
...     if node.left_subtree:
...         collect_words(node.left_subtree, words)
...     if node.right_subtree:
...         collect_words(node.right_subtree, words)
>>> 
>>> def test_all_words(root):
...     assert root.freq == sum(WORD_FREQ.values()), f'Expected {sum(WORD_FREQ.values())} but got {root.freq}'
...     words = set()
...     collect_words(root, words)
...     assert words == set(WORD_FREQ.keys()), f'Expected {set(WORD_FREQ.keys())} but got {words}'
>>> root = build_huffman_tree()
>>> test_all_words(root)
>>> 
>>> def check_rightmost(node):
...     if not node.right_subtree:
...         assert node.rightmost_word == node.word, f'Expected {node.word} but got {node.rightmost_word}'
...         return
...     assert node.rightmost_word == node.right_subtree.rightmost_word
...     check_rightmost(node.left_subtree)
...     check_rightmost(node.right_subtree)
>>> root = build_huffman_tree()
>>> check_rightmost(root)
```


Tests for confidence_interval_95 (points: 1)
```
>>> assert confidence_interval_95([1 for _ in range(20)]) == '1.000 +- 0.000'
>>> assert confidence_interval_95([0 for _ in range(10)] + [1 for _ in range(10)]) == '0.500 +- 0.219'
>>> (mean_bin, var_bin) = confidence_interval_95(binary_search).split('+-')
>>> (mean_huff, var_huff) = confidence_interval_95(huffman).split('+-')
>>> assert float(mean_bin) > float(mean_huff)
>>> assert float(var_bin) < 0.01
>>> assert float(var_huff) > 0.01
>>> assert float(mean_bin) > 9
>>> assert float(mean_huff) < 9
```


Tests for entropy (points: 1)
```
>>> assert abs(entropy - 7.935746) < 1e-05, 'Entropy is incorrect'
```


Tests for generate_secret_word (points: 2)
```
>>> def test_generate_secret_word():
...     for _ in range(1000):
...         word = generate_secret_word()
...         assert word in WORD_FREQ, f"Generated word '{word}' not in word frequency dictionary"
...     samples = 100000
...     counts = {}
...     for _ in range(samples):
...         word = generate_secret_word()
...         counts[word] = counts.get(word, 0) + 1
...     top_words = sorted(WORD_FREQ.items(), key=lambda x: x[1], reverse=True)[:10]
...     for (word, _) in top_words:
...         expected_freq = WORD_FREQ[word] / total_freq
...         actual_freq = counts.get(word, 0) / samples
...         assert abs(expected_freq - actual_freq) < 0.005, f"Word '{word}' frequency {actual_freq:.3f} differs too much from expected {expected_freq:.3f}"
>>> test_generate_secret_word()
```


Tests for huffman_game (points: 2)
```
>>> def test_huffman_game_public():
...     """tests for huffman_game."""
...     test_word = 'test'
...     (deduced, guesses) = huffman_game(create_game(test_word))
...     assert deduced == test_word, f'Expected {test_word} but got {deduced}'
...     assert isinstance(guesses, int), f'Expected guesses to be int but got {type(guesses)}'
...     assert guesses > 0, f'Expected positive number of guesses but got {guesses}'
...     assert guesses <= 15, f'Used more guesses than words'
...     for word in ['find', 'time', 'great', 'anyone', 'lying']:
...         (deduced, guesses) = huffman_game(create_game(word))
...         assert deduced == word, f'Expected {word} but got {deduced}'
...         assert isinstance(guesses, int)
...         assert guesses > 0
...         assert guesses <= 11
>>> test_huffman_game_public()
```
